# Analiza danych przestrzennych - ćwiczenia laboratoryjne 2023/2024

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania i oceny notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Jeżeli zestaw zadań wymaga skorzystania z fragmentów kodu opracowanego w ramach wcześniejszych zestawów zadań należy je umieścić we wskazanej komórce notatnika.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu, odpowiedzi uzupełniane manualnie nie podlegają ocenie.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają na celu wyłapać podstawowe błędy. Przejście przez kod wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie i zostanie ocenione na maksymalną liczbę punktów.
- Zadania należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zadania, które powodują wyświetlenie komunikatu o błędzie przerywającym wykonywanie kodu nie podlegają ocenie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem). Rozmiar poziomy figur nie powinien przekraczać 20 cali.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (nazwa zespołu oraz imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespołu:
Członkowie:

---

## Zestaw zadań 3: Elementy statystyki opisowej 3 (Notatnik 1/3)

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Zadanie 1: Przygotowanie danych [10 pkt]

#### a) Natężenie ruchu rowerowego 

Plik raw_bike_traffic_data.csv zawiera dane na temat ruchu rowerowego na moście Fremont Bridge w Seattle pobrane ze strony
https://data.seattle.gov/Transportation/Fremont-Bridge-Bicycle-Counter/65db-xm6k.

Zaimportuj zawarte w pliku dane do notatnika i utwórz na ich podstawie tablicę `bike_traffic_data` wykonaną według następujących wytycznych:

- Tablica powinna zawierać wyłącznie następujące elementy:</br>
    > Indeks: Date</br>
    > `pandas.PeriodIndex` zawierający okresy czasu, w których dokonywane były zliczenia rowerzystów.
    
    > Kolumny: Total, East, West</br>
    > Kolumny zawierają liczbę rowerzystów przejeżdżających przez most w poszczególnych godzinach w sumie oraz z podziałem na zachodnią i wschodnią stronę mostu.
- Pozostaw w tablicy wyłącznie wiersze dotyczące 2022 roku.

Zapisz tablicę do pliku bike_traffic_data.pkl.

Wyświetl przygotowaną tablicę w taki sposób aby widoczne było po 12 pierwszych i ostatnich wierszy, a zawarte w tablicy wartości wyświetlały się bez miejsc po przecinku.

<!-- grade_id: cell-aa85199b5cb633d2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
bike_traffic_data = pd.read_csv("raw_data/raw_bike_traffic_data.csv")
bike_traffic_data = bike_traffic_data.set_index(pd.PeriodIndex(data = bike_traffic_data["Date"], freq="H")).drop("Date", axis=1)
bike_traffic_data
bike_traffic_data.columns = ["Total", "East", "West"]
bike_traffic_data = bike_traffic_data[bike_traffic_data.index.year == 2022]
bike_traffic_data.to_pickle("bike_traffic_data.pkl")
bike_traffic_data_display = pd.concat([bike_traffic_data.head(12), bike_traffic_data.tail(12)]).astype(int)
bike_traffic_data_display
```

Output:
```
Total  East  West
Date                               
2022-01-01 00:00      0     0     0
2022-01-01 01:00      0     0     0
2022-01-01 02:00      1     0     1
2022-01-01 03:00      0     0     0
2022-01-01 04:00      0     0     0
2022-01-01 05:00      1     1     0
2022-01-01 06:00      1     1     0
2022-01-01 07:00      1     1     0
2022-01-01 08:00      3     3     0
2022-01-01 09:00      2     0     2
2022-01-01 10:00      0     0     0
2022-01-01 11:00      1     0     1
2022-12-31 12:00    124    63    61
2022-12-31 13:00    103    56    47
2022-12-31 14:00     98    46    52
2022-12-31 15:00     83    40    43
2022-12-31 16:00     68    28    40
2022-12-31 17:00     35    19    16
2022-12-31 18:00     28    15    13
2022-12-31 19:00     14     6     8
2022-12-31 20:00     20    13     7
2022-12-31 21:00     16     6    10
2022-12-31 22:00     23    15     8
2022-12-31 23:00     25    15    10
```

<!-- grade_id: cell-683b0a3c6c74d0c6, points: 2, grade: true, solution: false -->
```python
### Komórka testowa
# Test 1 Kształt tablicy
assert bike_traffic_data.shape == (8760, 3)
# Test 2 Kolumny
assert len(set(list(bike_traffic_data.columns)) & set(['Total', 'East', 'West'])) == 3
# Test 3 Indeks
assert (bike_traffic_data.index.name == 'Date') == True
assert isinstance(bike_traffic_data.index, pd.PeriodIndex)
assert bike_traffic_data.index.values.min() == pd.Period('2022-01-01 00:00', 'H')
assert bike_traffic_data.index.values.max() == pd.Period('2022-12-31 23:00', 'H')
# Test 4 Zawartość kolumn
assert np.all(np.isclose(bike_traffic_data[['Total', 'East', 'West']].min().values, np.array([0, 0, 0]))) == True
assert np.all(np.isclose(bike_traffic_data[['Total', 'East', 'West']].max().values, np.array([717, 464, 635]))) == True
```

#### b) Pogoda

Pliki 2022_01.txt, 2022_02.txt ...  2022_12.txt zawierają miesięczne raporty pogodowe (F6 Preliminary Local Climate Data Report) ze stacji NWS Seattle/Tacoma, WA pobrane ze strony https://forecast.weather.gov/product.php?site=SEW&product=CF6&issuedby=SEW.

Wyjaśnienie wykorzystanych w raportach symboli i skrótów można znaleźć pod następującym linkiem: https://www.weather.gov/grr/climateF6explain.

Zaimportuj zawarte w plikach dane do notatnika i utwórz na ich podstawie tablicę `weather_data` wykonaną według następujących wytycznych:

- Tablica powinna zawierać wyłącznie następujące elementy:</br>
    > Indeks: Date</br>
    > `pandas.PeriodIndex` zawierający okresy czasu, których dotyczą dane z raportów pogodowych.</br>
    
    > Kolumny: 2-poziomowe nagłówki w podanej niżej kolejności (nazwa 1 poziomu: podległe jej nazwy 2 poziomu) </br>
    >> Temperature: Min, Avg, Max</br>
    >> Precipitation: Total</br>
    >> Snow: Fall, Depth</br>
    >> Wind:  Avg speed, Max speed</br>
    >> Sky: Cloud cover</br>

- Zamień oznaczenia opadów śladowych na wartość 0.
- Przelicz występujące w tablicy wartości do *C (temperatura), mm (opady), m/s (wiatr) i ułamków (zachmurzenie).

Zapisz tablicę do pliku weather_data.pkl.

Wyświetl przygotowaną tablicę w taki sposób aby widoczne było po 10 pierwszych i ostatnich wierszy, a zawarte w tablicy liczby wyświetlały się z dokładnością 2 miejsc po przecinku.

<!-- grade_id: cell-186c2f35c4f9aca6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
days = [31, 28, 31, 30, 31, 30, 31, 31, 30, 31, 30, 31]
months = ['01', '02', '03', '04', '05', '06', '07', '08', '09', '10', '11', '12']
weather_data = pd.DataFrame()
for i in range(12) :
    x = pd.read_csv(f"raw_data/raw_weather_data/2022_{months[i]}.txt", header=13, sep="\s+")
    x = x[1:days[i]+1]
    weather_data = pd.concat([weather_data, x], ignore_index=True)
dates = pd.date_range(start='2022-01-01', end='2022-12-31')
weather_data.index = pd.PeriodIndex(dates, freq="D")
weather_data = weather_data.loc[:, ["MIN", "AVG", "MAX", "WTR", "SNW", "DPTH", "SPD", "SPD.1", "S-S"]]

weather_data.index.name = "Date"

weather_data["WTR"] = weather_data["WTR"].replace(["T"], 0.0)
weather_data["DPTH"] = weather_data["DPTH"].replace(["T"], 0.0)
weather_data["SNW"] = weather_data["SNW"].replace(["T"], 0.0)
weather_data = weather_data.replace("M", np.nan)
weather_data = weather_data.astype(float)

weather_data["MIN"] = (weather_data["MIN"] - 32)*5/9
weather_data["MAX"] = (weather_data["MAX"] - 32)*5/9
weather_data["AVG"] = (weather_data["AVG"] - 32)*5/9
weather_data["WTR"] = weather_data["WTR"] * 25.4
weather_data["SNW"] = weather_data["SNW"] * 25.4
weather_data["DPTH"] = weather_data["DPTH"] * 25.4
weather_data["SPD"] = weather_data["SPD"] *  1.60934 / 3.6
weather_data["SPD.1"] = weather_data["SPD.1"] *  1.60934 / 3.6
weather_data["S-S"] = weather_data["S-S"] / 10

weather_data.columns = pd.MultiIndex.from_tuples([('Temperature', 'Min'), ('Temperature', 'Avg'), ('Temperature', 'Max'),
    ('Precipitation', 'Total'),
    ('Snow', 'Fall'), ('Snow', 'Depth'),
    ('Wind', 'Avg speed'), ('Wind', 'Max speed'), ('Sky', 'Cloud cover')])

weather_data.to_pickle("weather_data.pkl")

pd.concat([weather_data.head(10), weather_data.tail(10)]).style.format(precision=2)
```

Output:
```
<>:6: SyntaxWarning: invalid escape sequence '\s'
<>:6: SyntaxWarning: invalid escape sequence '\s'
C:\Users\mikos\AppData\Local\Temp\ipykernel_3140\4118416122.py:6: SyntaxWarning: invalid escape sequence '\s'
  x = pd.read_csv(f"raw_data/raw_weather_data/2022_{months[i]}.txt", header=13, sep="\s+")
<pandas.io.formats.style.Styler at 0x26b905bda00>
```

<!-- grade_id: cell-e70d426962ca7e21, points: 8, grade: true, solution: false -->
```python
### Komórka testowa
# Test 1 Kształt tablicy
assert weather_data.shape == (365, 9)
# Test 2 Kolumny
assert len(set(list(weather_data.columns)) & set([('Temperature', 'Min'), ('Temperature', 'Avg'), ('Temperature', 'Max'), ('Precipitation', 'Total'),
            ('Snow', 'Fall'), ('Snow', 'Depth'), ('Wind', 'Avg speed'), ('Wind', 'Max speed'), ('Sky', 'Cloud cover')])) == 9
# Test 3 Indeks
assert (weather_data.index.name == 'Date') == True
assert isinstance(weather_data.index, pd.PeriodIndex) == True
assert weather_data.index.values.min() == pd.Period('2022-01-01', 'D')
assert weather_data.index.values.max() == pd.Period('2022-12-31', 'D')
# Test 4 Zawartość kolumn
assert np.all(np.isclose(weather_data.min().values, np.array([-7.77777778, -5.55555556, -3.33333333, 0., 0., 0., 0.89407778, 3.12927222, 0.]))) == True
assert np.all(np.isclose(weather_data.max().values, np.array([20., 26.66666667, 35., 75.438, 66.04, 76.2, 8.31492333, 16.54043889, 1.]))) == True
```